```python
# Initialize Otter
import otter
grader = otter.Notebook("sciPyPart1.ipynb")
```

## Lecture Section

In this lecture, we will cover the `scipy` library. Specifically, the `scipy.stats` module.
We will cover:
* Distributions (pdf, cdf), simulations (normal, binomial, etc.), and visualizations (plotting).
* Hypothesis testing
    * One-sample t-test
    * Two-sample t-test
    * Chi-square test of independence

### Simulations

Similar to `numpy`, we can simulate distributions with `scipy`, too. The underlying foundation of `SciPy` *is* `numpy`!

The code below simulates normal, uniform, and poisson distributions. It also shows some summary statistics of these distributions.

```python
import scipy.stats as stats
import numpy as np
import seaborn as sns
import matplotlib.pyplot as plt
import pandas as pd

normal_samples = stats.norm.rvs(loc=0, scale=1, size=1000)  # Normal Distribution
uniform_samples = stats.uniform.rvs(loc=0, scale=10, size=1000)  # Uniform distribution
poisson_samples = stats.poisson.rvs(mu=3, size=1000)  # Poisson distribution


print("Normal Distribution:")
print("Mean: ", np.mean(normal_samples))
print("Std Dev: ", np.std(normal_samples))

print("\nUniform Distribution:")
print("Mean: ", np.mean(uniform_samples))
print(f"Range: ", np.min(uniform_samples), " - ", np.max(uniform_samples))

print("\nPoisson Distribution:")
print("Mean: ", np.mean(poisson_samples))
print(f"Variance: ", np.var(poisson_samples))
```

Now, using matplotlib, we can plot each of these distributions. Make sure you understand the code below before moving on. Try removing setting the `kde` parameter to  `False` and see what happens!

```python
sns.histplot(normal_samples, kde=True)
plt.title("Simulated Normal Distribution")
plt.xlabel("Value")
plt.show()
```

```python
sns.histplot(uniform_samples, kde=False)
plt.title("Simulated Uniform Distribution")
plt.xlabel("Value")
plt.show()
```

We have to do a little extra work for the Poisson distribution, because Seaborn is trying to create bins for continuous data when we want one bar per integer value. We can directly tell it how many bins we want. Remember this when graphing discrete distributions!

```python
# Simulate Poisson Distribution
min_val, max_val = poisson_samples.min(), poisson_samples.max()
bins = np.arange(min_val - 0.5, max_val + 1.5, 1)  # bin edges centered on integers

sns.histplot(poisson_samples, bins=bins, kde=False)
plt.title("Simulated Poisson Distribution")
plt.xlabel("Counts")
plt.tight_layout()
plt.show()
```

### Working With Real Data

The section above was mostly review... let's try and use some real data. We will use the 1tips` dataset from `seaborn`.

https://www.kaggle.com/datasets/ranjeetjain3/seaborn-tips-dataset

```python
tips = sns.load_dataset("tips")
print(tips.head(24))
```

Let's say we want to estimate the probability of getting a tip over $5 a certain number of times. We can use what we've learned so far to do that!

Explaining the code below (by line number):
1. Creates a series of 0's and 1's. Each index corresponds to the same index in our dataset. If there is a 1, that means the tip at that particular index in our dataset is greater than 5. For example, the first tip greater than 5 in our dataset is at index 23, so index 23 is a 1 in our Series object.
2. Creates array of all our indexes so that we can plot in the future.
3. Adds the number of 1s we have in our Series object, effectively getting the number of tips greater than 5 from our dataset.
4. The total number of tips in our dataset. Don't let the variable name we are measuring fool you - it includes the 0s, too!


```python
high_tip = (tips["tip"] > 5).astype(int)
x = np.arange(0, len(high_tip) + 1)
successes = high_tip.sum()
trials = len(high_tip)
high_tip
```

```python
trials
```

```python
successes
```

Now we can get our probability distribution. Since we are working with successes/trials - this problem is clearly binomial. We use `stats.binom.pmf` to get our distribution. We set `n` to the number of trials, and `p` to our observed probability of success.

```python
pmf = stats.binom.pmf(x, n=trials, p=successes/trials)
plt.plot(x, pmf)
plt.title("Binomial Distribution of High Tips")
plt.xlabel("# Tips > 5")
plt.ylabel("Probability")
plt.xlim(0, 100) # limits our x-axis
plt.show()
```

We have the highest probability of getting 20 tips greater than 5 dollars! Of course... that's if you're the sole server in the entire restaurant and your quality of service remains unchanged, though :)

FOr fun, we can also obtain the cdf, or cumulative distribution. It works as expected!

```python
cdf = stats.binom.cdf(x, n=trials, p=successes/trials)
plt.plot(x, cdf)
plt.title("Binomial Cumulative Distribution of High Tips")
plt.xlabel("# Tips > 5")
plt.ylabel("Probability")
plt.xlim(0, 100) # limits our x-axis
plt.show()
```

Let's try one more problem!

This time, we want to model the distribution of the total bill. This follows a normal distribution, so we will use `stats.norm`. Like all normal distribution, we only need the mean and the standard deviation to create a model. Once we have those, we can call our `cdf`, `pdf`, and `ppf` (for discrete: `pmf`, too).
* `pmf` : probability mass function (discrete) - probability that a random variable takes on a particular value
* `pdf` : probability density function (continuous) - probability that a random variable falls between two numbers
* `cdf` : cumulative distribution function - probability that a random variable is less than or equal to a specific point (area under the `pdf` up to a certain point - integral of `pdf`)
* `ppf` : percentage point function/inverse cdf - provides the x values that a given percentage of observations fall below

```python
mean_bill = tips["total_bill"].mean()
std_bill = tips["total_bill"].std()
print("P(X < 20):", stats.norm.cdf(20, loc=mean_bill, scale=std_bill)) # probability of total bill being less than $20
print("x such that P(X < x) = 0.95:", stats.norm.ppf(0.95, loc=mean_bill, scale=std_bill)) # observation x = $34.429... and all prior make up 95% of the distribution
```

We can plot the distribution, too!

```python
x = np.linspace(0, 60, 300)
plt.plot(x, stats.norm.pdf(x, loc=mean_bill, scale=std_bill))
plt.title("Normal Distribution of Total Bill")
plt.xlabel("Total Bill")
plt.ylabel("Probability")
plt.show()
```

And the cdf...

```python
x = np.linspace(0, 60, 300)
plt.plot(x, stats.norm.cdf(x, loc=mean_bill, scale=std_bill))
plt.title("Normal Distribution of Total Bill")
plt.xlabel("Total Bill")
plt.ylabel("Probability")
plt.show()
```

### Hypothesis Testing

It is easy to do hypothesis testing in Python, too.

For example, maybe we hypothesize that the average total bill is around $20 dollars. To test this, we use `stats.ttest_1samp()`, or a 1-sample t-test. We give the column of data that we are interested in as the first argument, and we set `popmean` equal to our hypothesized value.

```python
t_stat, p_val = stats.ttest_1samp(tips["total_bill"], popmean=20)
print("One-sample t-test p-value:", p_val)

print("\n\nActual mean: ", np.mean(tips["total_bill"]))
```

We don't have strong evidence to reject our null hypothesis, so we cannot say that the population's total bill is much different from $20.


We can also do two-sample t-tests. Let's compare the total_bill column for lunch and dinner groups. We want to know if the mean total bill for each group is equal. For this, we use `stats.ttest_ind()`. We give the columns of data we are comparing as the first two arguments, then we set `equal_var` to `True` (default) or `False`. A good rule of thumb: if the ratio of the large sample variance to the smaller is less than 4, we can generally assume the population variances are the same (or similar enough).

```python
lunch = tips[tips["time"] == "Lunch"]["total_bill"]
dinner = tips[tips["time"] == "Dinner"]["total_bill"]
print(np.var(dinner)/np.var(lunch))
t_stat, p_val = stats.ttest_ind(lunch, dinner, equal_var=True)
print("Two-sample t-test (Lunch vs Dinner) p-value:", p_val)
```

Our p-value is very small... it is likely that in the population, the average total bill of lunch-havers differs from dinner-goers!


We can also perform a chi-square test of independence. This test aims to determine if two categorical variables are related to one-another. WHen working with real data, it's often easier to create the contingency table for the `stats/chi2_contingency()` function. This function returns a lot of information, including the p-value, the degrees of freedom, and the expected frequencies.

```python
contingency_table = pd.crosstab(tips["day"], tips["smoker"])
chi2, p, dof, expected = stats.chi2_contingency(contingency_table)
print("Chi-square test (Day vs Smoker) p-value:", p)
print('Expected: ', expected)
```

If we have data that is already separated, the function will create the contingency table on its own:

```python
data = [
    [7, 9, 8, 12, 5],  # Male: average work hours per day
    [8, 9, 10, 9, 6]   # Female: average work hours per day
]

chi2_stat, p_value, dof, expected = stats.chi2_contingency(data)

print("Chi-Square Statistic:", chi2_stat)
print("P-value:", p_value)
print('Expected: ', expected)
```

Finally, scipy lets us easily do permutation tests. Permutation tests involve shuffling data-points between or within groups to create 'new' resampled datasets. The goal is to see how much the test statistic (in the case below: difference in means) varies across the 'new' datasets. Then, we want to know how often a test static like the one we observed in our real data occurs.

For more information on this process, check out this page:
https://jillxoreilly.github.io/StatsCourseBook/permutation_intro.html

```python
smokers = tips[tips["smoker"] == "Yes"]["tip"]
non_smokers = tips[tips["smoker"] == "No"]["tip"]
def statistic(x, y):
    return np.mean(x) - np.mean(y)
perm_p = stats.permutation_test((smokers.values, non_smokers.values), statistic, n_resamples=10000, alternative='two-sided')
print("Permutation test (Smokers vs Non-Smokers tip) p-value:", perm_p.pvalue)
```

## Assignment Section

**Question 1.** We want to estimate the probability of our total bill being less than 8 dollars a certain number of times. Use the 'high_tip' example above to create the pmf graph.

```python
low_bill = ...
x = ...
successes = ...
trials = ...

pmf = ...
plt.plot(x, pmf)
plt.title("Binomial Distribution of Low Bills")
plt.xlabel("# Bills < 7")
plt.ylabel("Probability")
plt.xlim(0, 40) # limits our x-axis
plt.show()
```

```python
grader.check("q1")
```

**Question 2.** We hypothesize that the average tip is around $3 dollars. Test this with a one-sample t-test.

```python
t_stat, p_val = ...
print("One-sample t-test for tip vs $3:", p_val)
```

```python
grader.check("q2")
```

**Question 3.** Run a permutation test between lunch total_bill and dinner total_bill. Use 10,000 for n_resamples.

```python
lunch = ...
dinner = ...
def statistic(x, y):
    ...
perm_p = ...
print("Permutation test (Lunch vs Dinner Total Bill) p-value:", perm_p.pvalue)
```

```python
grader.check("q3")
```

**Question 4.** We want to model the distribution of the 'tips' by assuming it folllows a normal distribution. Then, find the probability that a randomly selected tip is greater than or equal to $5.

```python
mean_tip = ...
std_tip = ...
ans = ...
print("P(X >= 5):", ans) # tip being greater than 5
```

```python
grader.check("q4")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: 3)
```
>>> assert low_bill.to_list() == [0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
>>> assert x.tolist() == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50, 51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67, 68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84, 85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100, 101, 102, 103, 104, 105, 106, 107, 108, 109, 110, 111, 112, 113, 114, 115, 116, 117, 118, 119, 120, 121, 122, 123, 124, 125, 126, 127, 128, 129, 130, 131, 132, 133, 134, 135, 136, 137, 138, 139, 140, 141, 142, 143, 144, 145, 146, 147, 148, 149, 150, 151, 152, 153, 154, 155, 156, 157, 158, 159, 160, 161, 162, 163, 164, 165, 166, 167, 168, 169, 170, 171, 172, 173, 174, 175, 176, 177, 178, 179, 180, 181, 182, 183, 184, 185, 186, 187, 188, 189, 190, 191, 192, 193, 194, 195, 196, 197, 198, 199, 200, 201, 202, 203, 204, 205, 206, 207, 208, 209, 210, 211, 212, 213, 214, 215, 216, 217, 218, 219, 220, 221, 222, 223, 224, 225, 226, 227, 228, 229, 230, 231, 232, 233, 234, 235, 236, 237, 238, 239, 240, 241, 242, 243, 244]
```
```
>>> assert successes == 7
>>> assert trials == 244
```
```
>>> import numpy as np
>>> assert np.allclose(pmf, np.array([0.000823149654, 0.00593223464, 0.0212884623, 0.0507210339, 0.0902598989, 0.127963401, 0.150550471, 0.151185705, 0.132287492, 0.102456135, 0.0711140686, 0.0446815667, 0.0256243443, 0.0135066327, 0.00658234629, 0.00298103447, 0.00126017886, 0.000499192223, 0.000185939157, 6.53243939e-05, 2.17058904e-05, 6.83842398e-06, 2.04733023e-06, 5.83663764e-07, 1.58742589e-07, 4.1259677e-08, 1.02647006e-08, 2.44787202e-09, 5.60325135e-10, 1.23266638e-10, 2.60923053e-11, 5.32003176e-12, 1.04590814e-12, 1.98456423e-13, 3.6376289e-14, 6.44643096e-15, 1.10538309e-15, 1.83537208e-16, 2.95298039e-17, 4.60694333e-18, 6.97359038e-19, 1.02483144e-19, 1.46301534e-20, 2.02993199e-21, 2.73888983e-22, 3.59535467e-23, 4.59395437e-24, 5.71615114e-25, 6.92912485e-26, 8.18630784e-27, 9.42979764e-28, 1.0594577e-28, 1.1614146e-29, 1.24268865e-30, 1.29822998e-31, 1.32462284e-32, 1.320431e-33, 1.28631815e-34, 1.22493123e-35, 1.14057102e-36, 1.03870568e-37, 9.25401473e-39, 8.0675057e-40, 6.88366637e-41, 5.75000349e-42, 4.70302135e-43, 3.76734801e-44, 2.95617836e-45, 2.27271345e-46, 1.71221364e-47, 1.26429277e-48, 9.15141496e-50, 6.49458716e-51, 4.51967108e-52, 3.08475669e-53, 2.06518169e-54, 1.35637905e-55, 8.74076257e-57, 5.5273999e-58, 3.43045382e-59, 2.0897543e-60, 1.24969419e-61, 7.33714661e-63, 4.22974137e-64, 2.39447384e-65, 1.33125376e-66, 7.26957825e-68, 3.89939063e-69, 2.05477095e-70, 1.06376992e-71, 5.41111279e-73, 2.70467825e-74, 1.32852248e-75, 6.4132658e-77, 3.04283237e-78, 1.41904241e-79, 6.50519178e-81, 2.9315667e-82, 1.29879537e-83, 5.65728539e-85, 2.42284585e-86, 1.02027348e-87, 4.22476111e-89, 1.72029517e-90, 6.88871849e-92, 2.71285679e-93, 1.05071637e-94, 4.00249225e-96, 1.49960543e-97, 5.5263592e-99, 2.00322572e-100, 7.14268341e-102, 2.50521333e-103, 8.64350514e-105, 2.93363469e-106, 9.7949278e-108, 3.21723403e-109, 1.03957651e-110, 3.30466822e-112, 1.03347777e-113, 3.17965208e-115, 9.62422154e-117, 2.86589705e-118, 8.39585635e-120, 2.4197939e-121, 6.86118778e-123, 1.91392721e-124, 5.25234683e-126, 1.41800898e-127, 3.76614429e-129, 9.84013682e-131, 2.52920707e-132, 6.3949712e-134, 1.59057689e-135, 3.89154936e-137, 9.36550401e-139, 2.21701401e-140, 5.16203946e-142, 1.18215849e-143, 2.6626646e-145, 5.89830766e-147, 1.28496513e-148, 2.75289671e-150, 5.79967617e-152, 1.20146831e-153, 2.44733833e-155, 4.90146657e-157, 9.65127033e-159, 1.86828959e-160, 3.55532127e-162, 6.65060098e-164, 1.22281707e-165, 2.20978767e-167, 3.92461119e-169, 6.84962827e-171, 1.17470146e-172, 1.97944122e-174, 3.27699157e-176, 5.3295088e-178, 8.51408937e-180, 1.33593966e-181, 2.05868522e-183, 3.11532039e-185, 4.62890918e-187, 6.75257668e-189, 9.66997299e-191, 1.35923315e-192, 1.87508856e-194, 2.53835828e-196, 3.37154921e-198, 4.39330687e-200, 5.61535015e-202, 7.03916183e-204, 8.6528074e-206, 1.04283556e-207, 1.23204201e-209, 1.42663092e-211, 1.61881408e-213, 1.7997056e-215, 1.9599368e-217, 2.09041548e-219, 2.18315065e-221, 2.23204302e-223, 2.23353548e-225, 2.1870291e-227, 2.09499879e-229, 1.96278413e-231, 1.79807838e-233, 1.61018331e-235, 1.40913057e-237, 1.2047863e-239, 1.00605125e-241, 8.20246603e-244, 6.52739913e-246, 5.06825327e-248, 3.8383396e-250, 2.834217e-252, 2.03966012e-254, 1.43001359e-256, 9.76325461e-259, 6.48823883e-261, 4.19501135e-263, 2.63754526e-265, 1.61176925e-267, 9.56767138e-270, 5.51394049e-272, 3.08325295e-274, 1.67175296e-276, 8.78334989e-279, 4.46854455e-281, 2.19970548e-283, 1.04691324e-285, 4.8132505e-288, 2.13578992e-290, 9.13810508e-293, 3.76607216e-295, 1.49342211e-297, 5.69155018e-300, 2.08203264e-302, 7.30073291e-305, 2.45038025e-307, 7.85963078e-310, 2.40506435e-312, 7.00800192e-315, 1.940507e-317, 5.095e-320, 1.3e-322, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0]))
```


Tests for q2 (points: 2)
```
>>> assert np.isclose(t_stat, -0.019432641422916876)
```
```
>>> assert np.isclose(p_val, 0.9845119176410544)
```


Tests for q3 (points: 2)
```
>>> assert perm_p.pvalue < 0.005
```
```
>>> assert perm_p.pvalue > 0
```


Tests for q4 (points: 3)
```
>>> assert np.isclose(mean_tip, 2.99827868852459)
```
```
>>> assert np.isclose(std_tip, 1.3836381890011826)
```
```
>>> assert np.isclose(ans, 0.07398927292643875)
```
